---
title: "チュートリアルについて"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(gradethis)
library(learnr)
library(joholearnr)
library(learnrhash)

gradethis::gradethis_setup()

moodle_url <- "https://moodle.gs.chiba-u.jp/moodle/mod/assign/view.php?id=1444803"
```

## チュートリアルについて

### コードの書き方

本授業の宿題はチュートリアル形式になっています。チュートリアルでは、ブラウザ内で説明を読んだ後にRコードを書く場所が設けられているので、質問に従ってRコードを入力してください。

書き込むボックスの上に二つのボタンがあります：

- Run Code: コードを実行する
- Submit Answer: 答えを提出する

まずは **Run Code** をクリックして、自分のコードが思う通りの結果を返すかどうか確認してください。次に、**Submit Answer** をクリックして答えを提出してください。

**注意：**ブラウザによってクリックできないように見えるアイコンが出ることがありますが、無視して大丈夫です（クリックできます）：

![](https://raw.githubusercontent.com/joelnitta/joholearnr/refs/heads/main/inst/tutorials/00-learnr/arrow-bug.png){width=100px}

答えを提出したら、その答えが正しいかどうかが表示され、正しくない場合はヒントが表示されます（英語）。緑色の場合は正解、赤色の場合は不正解です。

答えの投稿は何回でも可能です。正解になるまで、コードを修正して提出してみてください。

例えば：

- Q1: 25/7はいくつですか？

```{r addition, exercise = TRUE, exercise.lines = 2}

```

```{r addition-check}
grade_result(
  pass_if(~ identical(.result, 25 / 7))
)
```

「Continue」ボタンを押して次の項目に進んでください。

### 他の質問の種類

コードを書く以外に、複数選択の質問もあります。

例えば、

- Q2: この授業で使うプログラミング言語の名前は？

```{r object-combine, echo=FALSE}
question("一つ選んでください：",
  answer("python", message = "これも人気のあるデータ解析用プログラミング言語ですが、この授業では使いません"),
  answer("S", message = "この授業で使っているプログラミング言語の一世代前です"),
  answer("R", correct = TRUE),
  answer("Arrrgh", message = "スペルが違います"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

## 確認と提出

全ての回答が終わったら、自分の名前と学籍番号を入力して、**Check Answers** をクリックし、提出内容を確認してください。

```{r, echo=FALSE, results = "asis"}
res <- knitr::knit_child(
  system.file("rmd", "summarize_score.Rmd", package = "joholearnr"),
  quiet = TRUE)
cat(res, sep = '\n')
```

これでは、答えはまだ提出できていません。

答えを提出するには、**Generate** ボタンをクリックしてください。

すると、とても長い文字列が上のボックスに現れます。この文字列には、あなたが答えた内容がすべて含まれています。この文字列をコピーして、Moodleの課題で提出してください。

そうすれば、完了です。

## 他の注意点

- 数分間操作しないとセッションがリセットされてしまいます（これまでの回答がすべて失われます）。宿題に取り組む際は、最後まで完了できる時間を確保しておきましょう。

- 左側のメニューの下に小さな「Start Over」ボタンがあります。これをクリックすると、セッションがリセットされます。

- ハッシュ（完成コード）は非常に長いため、コピー＆ペーストする際は全て正確に入力されていることを確認してください。
