---
title: "データの整理（1）"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(gradethis)
library(learnr)
library(joholearnr)
library(learnrhash)
suppressPackageStartupMessages(library(tidyverse))

gradethis::gradethis_setup()

tutorial_options(exercise.reveal_solution = FALSE)

moodle_url <- "https://moodle.gs.chiba-u.jp/moodle/mod/assign/view.php?id=1293876"
```

## 準備

必要なパッケージをロードしてから、今回の課題に使うデータを読み込みましょう。

今回のデータはURLから読み込みます。世界各国の経済データを含む「gapminder」というデータセットです。このデータセットについて詳しく知りたい場合は、<https://www.gapminder.org/> を参照してください。

まずは生データを読み込み、`gapminder_raw`として保存します。

```{r tidyverse-example, exercise = TRUE, exercise.lines = 6}
library(tidyverse)

gapminder_raw <- read_csv("https://raw.githubusercontent.com/swcarpentry/r-novice-gapminder/main/episodes/data/gapminder_data.csv")

gapminder_raw
```

### データの説明

`gapminder_raw`には以下の列があります：

- `country`: 国
- `year`: 年
- `pop`: 人口
- `continent`: 大陸
- `lifeExp`: 寿命（"Life Expectancy"）
- `gdpPercap`: 一人当たりの国内総生産（"GDP per capita"、米ドル単位）

## 列名を変更する

`gapminder_raw`には、大文字と小文字が混ざっている列名があります：`lifeExp`と`gdpPercap`。[授業で勉強したように](https://joelnitta.github.io/joho-shori/day4/#/%E3%83%87%E3%83%BC%E3%82%BF%E6%95%B4%E7%90%86%E3%81%AE%E9%96%A2%E6%95%B0)、列名は小文字に揃えた方が使いやすいので、列名を変更したいと思います。使い方の詳細については、[講義のスライド](https://joelnitta.github.io/joho-shori/day4/#/rename%E3%81%A7%E5%88%97%E5%90%8D%E3%82%92%E5%A4%89%E3%81%88%E3%82%8B)を参照してください。

- Q1: `rename()`または`clean_names()`関数を使って、この列名を`life_exp`と`gdp_percap`というふうに変更してください（`clean_names()`を使う場合は`janitor`パッケージをロードする必要があります）。`gapminder`というオブジェクトとして保存し、最後に`gapminder`の中身を表示してください。

```{r load-data-hide, echo = FALSE}
gapminder_raw <- read_csv("https://raw.githubusercontent.com/swcarpentry/r-novice-gapminder/main/episodes/data/gapminder_data.csv")
```

```{r rename-gapminder, exercise = TRUE, exercise.lines = 7, exercise.setup = "load-data-hide"}
# 列名を変更する
gapminder <- _____

# 保存したデータの中身を確認する
____
```

```{r clean-gapminder, echo = FALSE, exercise.setup = "load-data-hide"}
gapminder <- janitor::clean_names(gapminder_raw)
```

```{r rename-gapminder-check}
gapminder <- janitor::clean_names(gapminder_raw)
grade_this({
  fail_if_not_equal(
    x = last_user_code_call(.user_code),
    y = "gapminder",
    message = "最後に gapminder でデータの中身を表示する必要があります"
  )
  pass_if_equal(
    x = as.data.frame(.result),
    y = as.data.frame(gapminder)
  )
  fail()
})
```

### データの説明

列名を変更した結果、`gapminder`の列は次のようになります：

- `country`: 国
- `year`: 年
- `pop`: 人口
- `continent`: 大陸
- `life_exp`: 寿命（"Life Expectancy"）
- `gdp_percap`: 一人当たりの国内総生産（"GDP per capita"、米ドル単位）

## 行の順番を並び替える

`gapminder`のように、行が多いデータセットの特徴を知るために行の並び替えが大事です。`arrange()`は行を並び替える関数です。使い方の詳細については、[講義のスライド](https://joelnitta.github.io/joho-shori/day4/#/arrange%E3%81%A7%E8%A1%8C%E3%82%92%E4%B8%A6%E3%81%B3%E6%9B%BF%E3%81%88%E3%82%8B)を参照してください。

- Q2: `arrange()`関数を使って、`gapminder`データを寿命（短いから長い方へ）と並び替えてください。

```{r arrange-life-exp, exercise = TRUE, exercise.lines = 2, exercise.setup = "clean-gapminder"}
arrange(____, ____)
```

```{r arrange-life-exp-check}
  grade_this({
    pass_if_equal(
      x = as.data.frame(.result),
      y = as.data.frame(arrange(gapminder, life_exp))
    )
    fail_if_not_equal(
      x = inherits(.result, "data.frame"),
      y = TRUE,
      message = "最後の出力はデータフレームにする必要があります"
    )
    fail()
  })
```

`arrange()`はデフォルトで小さい方から大きい方へと並び替えますが、列名を`desc()`の中に入れると、逆の順になります（大きい → 小さい）。

- Q3: 一人当たりの国内総生産を大きな順から並び替えてください：

```{r arrange-desc-gdp, exercise = TRUE, exercise.lines = 2, exercise.setup = "clean-gapminder"}
arrange(____, desc(____))
```

```{r arrange-desc-gdp-check}
  grade_this({
    pass_if_equal(
      x = as.data.frame(.result),
      y = as.data.frame(arrange(gapminder, desc(gdp_percap)))
    )
    fail_if_not_equal(
      x = inherits(.result, "data.frame"),
      y = TRUE,
      message = "最後の出力はデータフレームにする必要があります"
    )
    fail()
  })
```

- Q4: このデータセットの中で、最も長い寿命の国はどこですか？ `arrange()`を使って求めてください。

```{r life-exp-highest-example, exercise = TRUE, exercise.lines = 2, exercise.setup = "clean-gapminder"}

```

```{r life-exp-highest, echo = FALSE}
question_text(
  text = "国名を言葉一つで記入してください：",
  answer_fn(function(value) {
    answer_lc <- str_to_lower(value)
    if (answer_lc == "japan") {
      correct()
    } else if (answer_lc == "日本") {
      correct()
    } else if (answer_lc == "にほん") {
      correct()
    } else if (answer_lc == "ニホン") {
      correct()
    } else if (str_detect(answer_lc, " |　")) {
      incorrect("国名のみを入力してください（スペースなど使用しないでください）")
    } else {
      incorrect()
    }
  }, label = "国名"),
  allow_retry = TRUE,
  trim = TRUE
)
```

## データを絞る

行の並び替えと同じように、条件で行を絞ることもデータの特徴を知る上で大事です。使い方の詳細については、[講義のスライド](https://joelnitta.github.io/joho-shori/day4/#/filter%E3%81%A7%E8%A1%8C%E3%82%92%E7%B5%9E%E3%82%8A%E8%BE%BC%E3%82%80)を参照してください。

- Q5: `filter()`関数を使って、`gapminder`から日本（"Japan"）のデータのみを抽出してください。
  - ヒント：「イコール」は`=`ではなく、`==`を使ってください！

```{r filter-japan, exercise = TRUE, exercise.lines = 2, exercise.setup = "clean-gapminder"}
filter(____, ____)
```

```{r filter-japan-check}
  grade_this({
    pass_if_equal(
      x = as.data.frame(.result),
      y = as.data.frame(filter(gapminder, country == "Japan"))
    )
    fail_if_not_equal(
      x = inherits(.result, "data.frame"),
      y = TRUE,
      message = "最後の出力はデータフレームにする必要があります"
    )
    fail()
  })
```

###

`filter()`に"or"（「もしくは」）という条件をかけたい時は、`|`を使います。例えば、大陸がアフリカ**または**アジアに絞る：

```{r filter-or-example, exercise = TRUE, exercise.lines = 6, exercise.setup = "clean-gapminder"}
filter(gapminder, continent == "Africa" | continent == "Asia")
```

- Q6: `filter()`関数を使って、`gapminder`から日本（"Japan"）、もしくは台湾（"Taiwan"）のデータのみを抽出してください。

```{r filter-japan-taiwan, exercise = TRUE, exercise.lines = 2, exercise.setup = "clean-gapminder"}
filter(____, ____)
```

```{r filter-japan-taiwan-check}
  grade_this({
    pass_if_equal(
      x = as.data.frame(.result),
      y = as.data.frame(filter(gapminder, country == "Japan"))
    )
    fail_if_not_equal(
      x = inherits(.result, "data.frame"),
      y = TRUE,
      message = "最後の出力はデータフレームにする必要があります"
    )
    fail()
  })
```

###

`filter()`では、複数の条件を`,`で区切ることで複数条件で絞り込みが可能です。たとえば、1997年に絞り、人あたりのGDPが3万米ドルを超える国を抽出します。

```{r filter-mult-example, exercise = TRUE, exercise.lines = 2, exercise.setup = "clean-gapminder"}
filter(gapminder, year == 1997, gdp_percap > 30000)
```

- Q7: 2007年に、アジア大陸で寿命が65歳未満（65歳を含まない）の国だけに絞ってください（結果は新しいオブジェクトとして保存する必要はありませんが、コードの出力はデータフレームにしてください）。

```{r filter-asia-life, exercise = TRUE, exercise.lines = 2, exercise.setup = "clean-gapminder"}
filter(____, ____, ____, ____)
```

```{r filter-asia-life-check}
  grade_this({
    pass_if_equal(
      x = as.data.frame(.result),
      y = as.data.frame(filter(gapminder, continent == "Asia", life_exp < 65, year == 2007))
    )
    fail_if_not_equal(
      x = inherits(.result, "data.frame"),
      y = TRUE,
      message = "最後の出力はデータフレームにする必要があります"
    )
    fail()
  })
```

## データを変更する

### 既存の列を変更する

列の中身を変える時は`mutate()`（「変化させる」）変数を使います。使い方の詳細については、[講義のスライド](https://joelnitta.github.io/joho-shori/day4/#/mutate%E3%81%A7%E5%88%97%E3%81%AE%E4%B8%AD%E8%BA%AB%E3%82%92%E5%A4%89%E3%81%88%E3%82%8B)を参照してください。

- Q8: `mutate()`関数を使用して、`pop`を万人単位に変換してください。

```{r mutate-pop, exercise = TRUE, exercise.lines = 2, exercise.setup = "clean-gapminder"}
mutate(____, ____)
```

```{r mutate-pop-check}
  grade_this({
    pass_if_equal(
      x = as.data.frame(.result),
      y = as.data.frame(mutate(gapminder, pop = pop/10000))
    )
    fail_if_not_equal(
      x = inherits(.result, "data.frame"),
      y = TRUE,
      message = "最後の出力はデータフレームにする必要があります"
    )
    fail()
  })
```

###

- Q9: `mutate()`関数を使用して、`gdp_percap`を米ドルから日本円に変換してください（1米ドル = 150日本円というレートを使ってください）。

```{r mutate-gdp-percap, exercise = TRUE, exercise.lines = 2, exercise.setup = "clean-gapminder"}
mutate(____, ____)
```

```{r mutate-gdp-percap-check}
  grade_this({
    pass_if_equal(
      x = as.data.frame(.result),
      y = as.data.frame(mutate(gapminder, gdp_percap = gdp_percap*150))
    )
    fail_if_not_equal(
      x = inherits(.result, "data.frame"),
      y = TRUE,
      message = "最後の出力はデータフレームにする必要があります"
    )
    fail()
  })
```

### 新しい列を作成する

`mutate(列名 = 新しい列を作成する計算式)`の形で使いますが、指定した`列名`が存在しない場合は、新しい列が作成されます。

たとえば、寿命の相対値を`life_exp_rel`として計算します。

```{r mutate-new-example, exercise = TRUE, exercise.lines = 3, exercise.setup = "clean-gapminder"}
mutate(gapminder, life_exp_rel = life_exp / max(life_exp))
```

- `max(life_exp)`は寿命のもっとも大きな値です。つまり、すべての寿命のデータをもっとも大きな寿命で割ることによって、相対的な値を作りました。

###

- Q10: 一人当たりの国内総生産（`gdp_percap`）を人口（`pop`）にかけることによって、国内総生産（`gdp`）を計算してください。

```{r mutate-gdp, exercise = TRUE, exercise.lines = 2, exercise.setup = "clean-gapminder"}
mutate(____, ____)
```

```{r mutate-gdp-check}
  grade_this({
    pass_if_equal(
      x = as.data.frame(.result),
      y = as.data.frame(mutate(gapminder, gdp = gdp_percap * pop))
    )
    fail_if_not_equal(
      x = inherits(.result, "data.frame"),
      y = TRUE,
      message = "最後の出力はデータフレームにする必要があります"
    )
    fail()
  })
```

## 確認と提出

```{r, echo=FALSE, results = "asis"}
res <- knitr::knit_child(
  system.file("rmd", "summarize_score.Rmd", package = "joholearnr"),
  quiet = TRUE)
cat(res, sep = '\n')
```
