---
title: "R入門（２）"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(gradethis)
library(learnr)
library(joholearnr)
library(learnrhash)

gradethis::gradethis_setup()
# do this just to maintain dep on {praise} so it will be installed in the
# shiny app
invisible(praise::praise())

tutorial_options(exercise.reveal_solution = FALSE)

# Need for successfull installing of {praise} in Q13
options(repos = c(CRAN = "https://cloud.r-project.org"))

moodle_url <- "https://moodle.gs.chiba-u.jp/moodle/mod/assign/view.php?id=1267594"
```

## ベクトルについて

### ベクトルの作り方

今までのオブジェクトは一つだけの値が入っていました。しかし、複数の値を入れることもできます。

複数の値を入れるには、`c()`を使います。

下のコードを実行してみましょう：

```{r concat-example, exercise = TRUE, exercise.lines = 2}
x <- c(2, 4, 6)
x
```

`x`には三つの値（2と4と6）が入りました。

Rでは、このようなオブジェクトを**「ベクトル」**（vector）と呼びます。

ベクトルのそれぞれの値を**「要素」**（element）と呼びます。

`x`の2番目の要素は`4`です。

###

- Q1: `c()`を使って、45, 55, 33という値を`length_cm`に入れてください。

```{r make-length-cm, exercise = TRUE, exercise.lines = 2}
length_cm <- ____
```

```{r make-length-cm-solution}
length_cm <- c(45, 55, 33)
```

```{r make-length-cm-check}
grade_this_code()
```

### ベクトルの計算

ベクトルの便利なところは、一気にベクトルの要素を全て同じように計算できるところです。

例えば、`x`のそれぞれの要素を2で割る：

```{r vector-calc-example, exercise = TRUE, exercise.lines = 2}
x <- c(2, 4, 6)
x / 2
```

###

- Q2: `length_cm`をインチに変換してください（新しいオブジェクトとして保存してもしなくてもいいですが、コードの出力はインチの値にしないといけません）。

```{r cm-convert, exercise = TRUE, exercise.lines = 2, exercise.setup = "make-length-cm-solution"}

```

```{r cm-convert-check}
grade_result(
  pass_if(~ identical(.result, c(45, 55, 33) / 2.54))
)
```

### ベクトルの種類

今までは、数字しか扱ってきませんでしたが、Rには他のベクトルの種類もあります。Rでは、ベクトルの種類を**「型」**（type）と呼びます。他の種類のベクトルも作り方は数字ベクトルと同じです：`c()`を使います。

例えば、文字ベクトル：

```{r character-example, exercise = TRUE, exercise.lines = 2}
fruits <- c("banana", "apple", "kiwi")
fruits
```

**注意：** 文字ベクトルを作るとき、データ（値）を**引用符**（`"`）で囲まないといけません。そうしないと、Rはそれが他のオブジェクトの名前だと思ってしまい、エラーが出ます。

```{r character-error-example, exercise = TRUE, error = TRUE}
fruits <- c(banana, apple, kiwi)
```

###

- Q3: 文字ベクトルを作ってください（新しいオブジェクトとして保存してもしなくてもいいですが、コードの出力は文字ベクトルにしないといけません）。

```{r make-char-vec, exercise = TRUE, exercise.lines = 2}

```

```{r make-char-vec-check}
grade_result(
  pass_if(~ identical(typeof(.result), "character"))
)
```

### 論理ベクトル

もう一つのベクトルの種類があります：**論理ベクトル**

例えば、コインを3回投げた結果を論理ベクトルとして保存しましょう：

```{r logical-example, exercise = TRUE, exercise.lines = 2}
coin_flip_results <- c(TRUE, FALSE, TRUE)
coin_flip_results
```

文字ベクトルと違って、論理ベクトルに入れられる値は`TRUE`か`FALSE`しかなく、引用符を使いません。

###

- Q4: 論理ベクトルを作ってください（新しいオブジェクトとして保存してもしなくてもいいですが、コードの出力は論理ベクトルにしないといけません）。

```{r make-log-vec, exercise = TRUE, exercise.lines = 2}

```

```{r make-log-vec-check}
grade_result(
  pass_if(~ identical(typeof(.result), "logical"))
)
```

### ベクトルの種類の確認

ベクトルの種類（型）を確認するには、`typeof()`を使います。

例えば：

```{r check-type-example, exercise = TRUE, exercise.lines = 2}
x <- c(2, 4, 6)
typeof(x)
```

（少し分かりづらいですが、Rでは"double"は「数字」という意味になります）

###

- Q5: `typeof()`を使って、`fruits`の種類を確認してください

```{r fruits-prep, include = FALSE}
fruits <- c("banana", "apple", "kiwi")
```

```{r fruits-type, exercise = TRUE, exercise.lines = 2, exercise.setup = "fruits-prep"}

```

```{r fruits-type-check}
grade_result(
  pass_if(~ identical(.result, typeof(fruits)))
)
```

- Q6: `typeof()`を使って、`length_cm`の種類を確認してください

```{r length-cm-type, exercise = TRUE, exercise.lines = 2, exercise.setup = "make-length-cm-solution"}

```

```{r length-cm-type-check}
grade_result(
  pass_if(~ identical(.result, typeof(length_cm)))
)
```

### データ型の自動変換

Rでは、あるベクトルには一つの型しか持てないというルールがあります。

そのため、異なる型のデータを結合させようとすると、自動的に同じ型に変換されます。

例えば、

```{r combine-vec-types-example, exercise = TRUE, exercise.lines = 2}
c(1, 2, "apple", "banana")
```

引用符`"`が付いているというのは、文字であることを示します。つまり、上のコードの出力の`"1"`は数字ではなく、文字です。

###

- Q7: 以下のコードを実行してください。結果の型はどれでしょうか？

```{r combine-type-example, exercise = TRUE, exercise.lines = 2}
c(TRUE, FALSE, 1, 0)
```

```{r object-combine, echo=FALSE}
question("一つ選んでください：",
  answer("論理", message = "論理ベクトルはTRUEかFALSE"),
  answer("文字", message = "文字ベクトルの要素は引用符が付いています"),
  answer("数字", correct = TRUE),
  answer("句読点", message = "Rに「句読点」というデータ型は存在しません"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

## 関数について

今まで使ってきた`c()`や`typeof()`はRでは**「関数」**（function）というものです。

関数は何かの入力を受け取り、それに対して何かの処理を行い、出力を返します。

- `c()`は入力した値を統合させて、一つのベクトルを返す関数です。
- `typeof()`は入力したオブジェクトの種類を返す関数です。

関数の書き方は、関数名をまず書いて、その後に括弧の中に入力を書きます：`関数名(入力)`

早速、新しい関数を使ってみましょう。まず、準備として`length_in`（インチ単位の長さのベクトル）を作りましょう：

```{r length-in-example, exercise = TRUE, exercise.lines = 2, exercise.setup = "make-length-cm-solution"}
length_in <- length_cm / 2.54
length_in
```

`length_in`の少数を切り捨てたいとしましょう。Rでは、少数を切り捨てる関数は`round()`です。

- Q8: `round()`を使って、`length_in`の少数を切り捨ててください。

```{r length-in-round-prep, include = FALSE, exercise.setup = "make-length-cm-solution"}
length_in <- length_cm / 2.54
```

```{r length-in-round, exercise = TRUE, exercise.lines = 2, exercise.setup = "length-in-round-prep"}

```

```{r length-in-round-solution}
round(length_in)
```

```{r length-in-round-check}
grade_this_code()
```

### 引数の使い方

関数の入力および他の設定を**「引数」**（argument）と言います。

`round()`に入力する数字の引数の名前は`x`です：

```{r arguments-1-example, exercise = TRUE, exercise.lines = 2}
round(x = 3.33333)
```

そして、`round()`にはもう一つの引数があります：`digits`。`digits`は残す少数の数を設定する引数です：

```{r arguments-2-example, exercise = TRUE, exercise.lines = 2}
round(x = 3.33333, digits = 2)
```

複数の引数がある場合、コンマで区切ります。

しかし、`round()`を最初に使った時は`x =`というふうに指定しなかったけど、ちゃんと切り捨てることができましたよね。なぜでしょう？

Rでは、引数の位置も決まっています。

1. `round()`の**1つ目**の引数は**切り捨てる数字**です（`x`）
2. `round()`の**2つ目**の引数は**残す少数の数**です（`digits`）

したがって、Rでは引数名を指定しない場合、**位置で引数を認識します**。

###

- Q9: `round()`の引数名`x`と`digits`を使って、`length_in`の少数を一つ残して切り捨ててください。

```{r round-use-args-names, exercise = TRUE, exercise.lines = 2, exercise.setup = "length-in-round-prep"}

```

```{r round-use-args-names-solution}
# order last ----
round(x = length_in, digits = 1)

# order first ----
round(digits = 1, x = length_in)

```

```{r round-use-args-names-check}
grade_this_code()
```

- Q10: **引数名を使わないで**、`length_in`の少数を一つ残して、`round()`で切り捨ててください。

```{r round-no-args-names, exercise = TRUE, exercise.lines = 2, exercise.setup = "length-in-round-prep"}

```

```{r round-no-args-names-solution}
round(length_in, 1)
```

```{r round-no-args-names-check}
grade_this_code()
```

###

`seq()`は数字ベクトルを作る関数です。`seq()`の引数は以下の通りです：

- `from`: 始まる数
- `to`: 終わる数
- `by`: 連続する数字の間隔

- Q11: `seq()`を使って、1から99まで（1と99を含めて）の奇数の数字ベクトルを作ってください

```{r seq-args, exercise = TRUE, exercise.lines = 2}

```

```{r seq-args-check}
grade_result(
  pass_if(~ identical(.result, seq(from = 1, to = 99, by = 2)))
)
```

### 関数を調べる

Rには非常にたくさんの関数があります。どんなにRをマスターしても、全部を覚えることはほとんどありませんし、必要でもありません。

ある関数についてその使い方（引数）を確認したいときは、`?関数名`、あるいは`help(関数名)`を実行します。

- Q12: `toupper()`の引数の名前はどれでしょうか？ `?` あるいは`help()`を使って、RStudioあるいは下のコンソールで確認してください。

```{r help-example, exercise = TRUE, exercise.lines = 2}
help(____)
```

```{r argument-name, echo=FALSE}
question("一つ選んでください：",
  answer("word", message = "もう一度ヘルプファイルを読んでください"),
  answer("x", correct = TRUE),
  answer("y", message = "もう一度ヘルプファイルを読んでください"),
  answer("character", message = "これは入力の種類です"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

## パッケージについて

Rには備え付けの関数がたくさんあり、この関数だけによるRの使い方を**「Base R」**（基本R）と呼びます。

しかし、やりたいことによってはBase Rでは足りない時もあります。Base Rの機能を拡張するものを**「パッケージ」**（Package）と呼びます。パッケージとは、何かの共通テーマを持った関数のまとまりです。以下はパッケージの例です：

- 図表を書くパッケージ`ggplot2`
- エクセルファイルを読み込むパッケージ`readxl`
- Rから電子メールを送るパッケージ`gmailr`

要するに、スマートフォンのアプリと同じようなものです。現在、Rには2万以上のパッケージがあり、情報処理に関することならほぼ何でもできます。

### パッケージのインストール

パッケージを使いたい場合、まずはそのパッケージをダウンロードする必要があります。

パッケージを自分のパソコンにダウンロード（インストール）する関数は`install.packages()`です。`()`の中にパッケージ名（文字）を入れます。

例えば、図表を書くパッケージ、`ggplot2`をインストールするコマンドは：

```{r install-package-example, eval = FALSE}
install.packages("ggplot2")
```

パッケージを一度インストールしたら、次にRを使う時は再度`install.packages()`をしなくていいです（すでにパソコンにダウンロードしてあるから）。

- Q13: `praise`パッケージをインストールしてください（出力が出るまで少し時間がかかるかもしれません、注意「Warning」が出ても気にしなくて大丈夫です）

```{r praise-install, exercise = TRUE, exercise.lines = 2}

```

```{r praise-install-solution}
install.packages("praise")
```

```{r praise-install-check}
grade_this_code()
```

### パッケージのロード

パッケージをインストール（ダウンロード）するのと、そのパッケージを使える状態にするのは別のことです。

パッケージを使える状態にする関数は`library()`です（自分のパソコンが図書館だとしたら、それぞれのパッケージが本になる、みたいなイメージです）。

`install.packages()`と二つの大きな違いがあります：

1. `install.packages()`は一つのパッケージについて一回だけすればいいですが、**`library()`はRを使う度に必要です。**
2. `install.packages()`を使う時はパッケージ名を引用符で囲みますが、**`library()`では引用符でパッケージ名を書きません。**

例えば、

```{r library-example, exercise = TRUE, exercise.lines = 2}
library(stats)
```

`library()`を実行した後、Rは特に何もメッセージを返しません。これはつまり、そのパッケージを使える状態に成功したという意味です（うまく行かなかったら、エラーが出ます）。

- Q14: `praise`パッケージをロードしてください。

```{r praise-load, exercise = TRUE, exercise.lines = 2}

```

```{r praise-load-solution}
library(praise)
```

```{r praise-load-check}
grade_this_code()
```

- Q15: `praise()`関数を使ってみてください

```{r praise-function, exercise = TRUE, exercise.lines = 2, exercise.setup = "praise-load-solution"}

```

```{r praise-function-solution}
praise()
```

```{r praise-function-check}
grade_this_code()
```

- Q16: `praise()`関数の出力はどんな型ですか？

```{r praise-output-type, echo = FALSE}
question("一つ選んでください：",
  answer("論理", message = "論理ベクトルはTRUEかFALSEです"),
  answer("文字", correct = TRUE),
  answer("数字", message = "数字には引用符が付いていません"),
  answer("句読点", message = "Rに「句読点」というデータ型は存在しません"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

## 確認と提出

```{r, echo=FALSE, results = "asis"}
res <- knitr::knit_child(
  system.file("rmd", "summarize_score.Rmd", package = "joholearnr"),
  quiet = TRUE)
cat(res, sep = '\n')
```
