---
title: "データの読み込み"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(gradethis)
library(learnr)
library(joholearnr)
library(learnrhash)
suppressPackageStartupMessages(library(tidyverse))

gradethis::gradethis_setup()

tutorial_options(exercise.reveal_solution = FALSE)

moodle_url <- "https://moodle.gs.chiba-u.jp/moodle/mod/assign/view.php?id=1444812"
```

## データをファイルから読み込む

### 準備

本授業では`tidyverse`パッケージ群に含まれているデータ読み込み関数を使用するため、まずは`tidyverse`をロードしましょう（チュートリアルではRStudioと違って、パッケージが立ち上がった時のメッセージは表示されませんが、「Run Code」を押してエラーが表示されなければ、正しくロードされています）。

```{r tidyverse-example, exercise = TRUE, exercise.lines = 2}
library(tidyverse)
```

### 練習ファイルのパス

`tidyverse`には練習用のデータが含まれています。練習用データの場所（パス）を表示するために、`readr_example()`関数を使います。例えば、

```{r readr-path-example, exercise = TRUE, exercise.lines = 2}
readr_example("challenge.csv")
```

- Q1: `readr_example()`を使って、`chickens.csv`のパスを表示してください。

```{r readr-path, exercise = TRUE, exercise.lines = 2}

```

```{r readr-path-solution}
readr_example("chickens.csv")
```

```{r readr-path-check}
grade_this_code()
```

- Q2: `chickens.csv`が入っているフォルダーの名前は何でしょうか？

```{r chickens-path, echo=FALSE}
question("一つを選んでください：",
  answer("Desktop"),
  answer("Library"),
  answer("Z:"),
  answer("extdata", correct = TRUE),
  answer("readr"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

### データの読み込み

`.csv`ファイルを読み込むために`read_csv()`関数を使います。`read_csv()`の最初の引数`file`は`.csv`ファイルのパスです。

- Q3: `read_csv()`と`readr_example()`を使って、`chickens.csv`を読み込み、`chickens`というオブジェクトとして保存してください。最後に`chickens`の中身を表示してください。

```{r chickens, exercise = TRUE, exercise.lines = 6}
# データを読み込んで、保存する
chickens <- _____

# 保存したデータの中身を確認する
____
```

```{r chickens-check}
chickens_df <- readr::read_csv(readr::readr_example("chickens.csv"))
grade_this({
  fail_if_not_equal(
    x = last_user_code_call(.user_code),
    y = "chickens",
    message = "最後に chickens でデータの中身を表示する必要があります"
  )
  pass_if_equal(
    x = as.data.frame(.result),
    y = as.data.frame(chickens_df)
  )
  fail()
})
```

- Q4: `chickens`にはいくつの行が入っていますか？

```{r chickens-rows, echo=FALSE}
question("一つを選んでください：",
  answer("4", message = "これは列の数です"),
  answer("2"),
  answer("5", correct = TRUE),
  answer("10"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

- Q5: `chickens`にはいくつの列が入っていますか？

```{r chickens-cols, echo=FALSE}
question("一つを選んでください：",
  answer("5", message = "これは行の数です"),
  answer("2"),
  answer("4", correct = TRUE),
  answer("10"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

- Q6: `chickens`の列のうち、数字データはどれでしょうか？

```{r chickens-cols-question, echo=FALSE}
question("一つを選んでください：",
  answer("chicken"),
  answer("sex"),
  answer("eggs_laid", correct = TRUE),
  answer("motto"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

## データをURLから読み込む

URLからでもファイルを読み込むことができます。ファイルのパスの代わりにURLを使用します。

- Q7:  <https://pos.it/r4ds-01-sales>というURLからCSVファイルを読み込み、`sales`というデータフレームに保存してください。最後の行に`sales`と打ってその中身を確認してください。

```{r read-csv-url, exercise = TRUE, exercise.lines = 3}

```

```{r read-csv-url-check}
sales_df <- readr::read_csv("https://pos.it/r4ds-01-sales")
grade_this({
  fail_if_not_equal(
    x = last_user_code_call(.user_code),
    y = "sales",
    message = "最後に sales でデータの中身を表示する必要があります"
  )
  pass_if_equal(
    x = as.data.frame(.result),
    y = as.data.frame(sales_df)
  )
  fail()
})
```

- Q8: `sales`の列のうち、文字データはどれでしょうか？

```{r sales-cols-question, echo=FALSE}
question("一つを選んでください：",
  answer("month", correct = TRUE),
  answer("year"),
  answer("brand"),
  answer("item"),
  answer("n"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

- Q9: `sales`にはいくつの行が入っていますか？

```{r sales-rows, echo=FALSE}
question("一つを選んでください：",
  answer("5", message = "これは列の数です"),
  answer("11"),
  answer("7", correct = TRUE),
  answer("4"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

- Q10: `sales`にはいくつの列が入っていますか？

```{r sales-cols, echo=FALSE}
question("一つを選んでください：",
  answer("7", message = "これは行の数です"),
  answer("3"),
  answer("5", correct = TRUE),
  answer("9"),
  random_answer_order = TRUE,
  allow_retry = TRUE
)
```

## Rの作業ディレクトリ

- Q11: 現在のRの作業ディレクトリをRの関数で確認してください。  
  - ヒント：関数名は"get working directory"の略です。

```{r r-wd, exercise = TRUE, exercise.lines = 2}

```

```{r r-wd-solution}
getwd()
```

```{r r-wd-check}
grade_this_code()
```

## 確認と提出

```{r, echo=FALSE, results = "asis"}
res <- knitr::knit_child(
  system.file("rmd", "summarize_score.Rmd", package = "joholearnr"),
  quiet = TRUE)
cat(res, sep = '\n')
```
